# Module 1 Hands-On Assignment

You are organizing your family contact list, but find that the current format of the phone numbers is flawed! The contact phone numbers are in different formats and you want valid phone numbers to all follow the same format.

Below, you can see that we have a file where each line consists of a person's name and their phone number.

First, run the code below to extract the lines.

```python
text = None

with open('assets/phonebook.txt', 'r') as file:
    lines = file.readlines()

print(*lines[:5],sep='\n')  # print the first 5 rows in the file
```

Output:
```
David Smith: +1-361-789-5031

Michael Rodriguez: 801 998 5951

Sarah Johnson: (270) 573-8188

Jane Williams: 571.358.9802

Laura Davis: (965) 511-7358
```

# Task 1: Create reg expression

These are possible formats a phone number could be stored in:

Parentheses: (409) 339-9584

Dashes: 943-641-7270

Dots: 574.874.7163

Spaces: 175 335 3625

International: +1-536-234-1765

Plain: 6565865784

For this task, you must define the 'reg' expression such that the given code is able to extract phone numbers in every one of these formats.

When we use reg.findall, we want to obtain a list of 3 strings containing the digits of each phone number separated by section. For example, if we have:

`str = John Smith: 123 456 7890`

We want `re.findall(reg, str)` on this line using our regex expression to obtain a list that contains the following:

['123','456','7890']

Look up how to use regex 'capturing groups' to accomplish this. The demo notebook associated with this assignment also contains hints on how you can accomplish this task. 

Assign your reg expression to `reg` in `task_1_solution`.

We use your reg expression in `extract_contents` to extract the contents of `lines` into a list where each entry in the list formatted like so:

`['First Last','000', '000', '0000']`

<!-- grade_id: cell-cbc369c659071340, points: 0, grade: true, solution: false -->
```python
import re


def extract_contents(reg):
    extracted_contents = []
    for line in lines:

        name, number = line.strip().split(':')
        number = number.strip()

        matches = reg.findall(number)

        for match in matches:
            parts = [part for part in match if part]
            extracted_contents.append([name] + parts)
    return extracted_contents
```

<!-- grade_id: cell-f8d3e7f2226f7599, points: 0, grade: false, solution: true -->
```python

def task_1_solution():

   
    pattern = r"""
    \b
    (?:\+1[-\s])?
    (?:\()?
    (\d{3})
    (?:\))?
    [-.\s]?
    (\d{3})
    [-.\s]?
    (\d{4})
    \b
    """


    reg = re.compile(pattern, re.VERBOSE)

    return extract_contents(reg), reg

```

```python
lines[21]
```

Output:
```
'Michael Williams: 556 541 8853\n'
```

```python
# Use the following cell to explore your solution

task_1_solution()[0][21]
```

Output:
```
['Sarah Wilson', '904', '357', '9522']
```

<!-- grade_id: cell-1389e58b8d461fd7, points: 60, grade: true, solution: false -->
```python

print(f"AG tests")
stu_contents, stu_reg = task_1_solution()


print(f"Your answer:\n{stu_contents, stu_reg }")

assert isinstance(
    stu_contents, list
), f"Your function should return a list "

def test_format(testLine):

    name, number = testLine.strip().split(':')
    number = number.strip()
    number_components = stu_reg.findall(number)
    parts = None
    for component in number_components:
        # Filter out empty strings and None values from the match
        parts = [part for part in component if part]
    return [name] + parts



assert (
    test_format("Jane Doe: (123) 456-7890") == ['Jane Doe', '123', '456', '7890']
), f"Your reg expression does not read entries in the format '(123) 456-7890' correctly "

assert (
    test_format("Jane Doe: 123-456-7890") == ['Jane Doe', '123', '456', '7890']
), f"Your reg expression does not read entries in the format '123-456-7890' correctly "

assert (
    test_format("Jane Doe: 123.456.7890") == ['Jane Doe', '123', '456', '7890']
), f"Your reg expression does not read entries in the format '123.456.7890' correctly "

assert (
    test_format("Jane Doe: 123 456 7890") == ['Jane Doe', '123', '456', '7890']
), f"Your reg expression does not read entries in the format '123 456 7890' correctly "

assert (
    test_format("Jane Doe: +1-123-456-7890") == ['Jane Doe', '123', '456', '7890']
), f"Your reg expression does not read entries in the format '1-123-456-7890' correctly "

assert (
    test_format("Jane Doe: 1234567890") == ['Jane Doe', '123', '456', '7890']
), f"Your reg expression does not read entries in the format '1234567890' correctly "

# Hidden tests are located in this cell
```

Output:
```
AG tests
Your answer:
([['David Smith', '361', '789', '5031'], ['Michael Rodriguez', '801', '998', '5951'], ['Sarah Johnson', '270', '573', '8188'], ['Jane Williams', '571', '358', '9802'], ['Laura Davis', '965', '511', '7358'], ['Jane Brown', '304', '948', '4696'], ['Chris Smith', '939', '455', '7006'], ['Sarah Garcia', '973', '516', '1118'], ['Michael Smith', '277', '961', '9814'], ['David Williams', '850', '853', '9220'], ['Sarah Jones', '129', '611', '3195'], ['Chris Garcia', '859', '449', '6647'], ['David Johnson', '227', '507', '2016'], ['Sarah Garcia', '613', '269', '9243'], ['Sarah Smith', '906', '714', '2381'], ['David Rodriguez', '193', '406', '1153'], ['Michael Williams', '556', '541', '8853'], ['Jane Johnson', '290', '616', '6165'], ['Alex Rodriguez', '810', '219', '7286'], ['John Miller', '673', '805', '8394'], ['Laura Brown', '843', '583', '8123'], ['Sarah Wilson', '904', '357', '9522'], ['David Garcia', '639', '976', '2570'], ['Jane Wilson', '465', '394', '2734'], ['Mich
```

# Task 2: Normalizing Data

Now, we'd like to normalize the phone numbers to the format: (123) 456-7890.

The function below will take in the ```stu_contents``` that you helped define in task 1 as an argument. It will then alter every entry in the ```stu_contents``` list to be a string in the (123) 456-7890 format. The function will return the ```extracted_contents_normalized``` list as a list of strings.

For example, 

The first entry in your task 1 `extracted_contents` is 

['David Smith', '838', '177', '8461']

This should be turned into

'David Smith: (838) 177-8461'

<!-- grade_id: cell-cec3113f5d8484ec, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-88f3262a40b3e624, points: 0, grade: false, solution: true -->
```python

def task_2_solution():

    extracted_contents, _ = task_1_solution()

    extracted_contents_normalized = []
    
    for i in extracted_contents:

        extracted_contents_normalized.append(i[0] + ': (' + i[1] + ') ' + i[2] + '-' + i[3])

    return extracted_contents_normalized
```

```python
# Use the following cell to explore your solution

task_2_solution()[0]
```

Output:
```
'David Smith: (361) 789-5031'
```

<!-- grade_id: cell-941c19b10c9e2a03, points: 40, grade: true, solution: false -->
```python

print(f"AG tests")
stu_ans = task_2_solution()


print(f"Your answer:\n{stu_ans}")

assert isinstance(
    stu_ans, list
), f"Your function should return a list"

# Hidden tests are located in this cell

del stu_ans

```

Output:
```
AG tests
Your answer:
['David Smith: (361) 789-5031', 'Michael Rodriguez: (801) 998-5951', 'Sarah Johnson: (270) 573-8188', 'Jane Williams: (571) 358-9802', 'Laura Davis: (965) 511-7358', 'Jane Brown: (304) 948-4696', 'Chris Smith: (939) 455-7006', 'Sarah Garcia: (973) 516-1118', 'Michael Smith: (277) 961-9814', 'David Williams: (850) 853-9220', 'Sarah Jones: (129) 611-3195', 'Chris Garcia: (859) 449-6647', 'David Johnson: (227) 507-2016', 'Sarah Garcia: (613) 269-9243', 'Sarah Smith: (906) 714-2381', 'David Rodriguez: (193) 406-1153', 'Michael Williams: (556) 541-8853', 'Jane Johnson: (290) 616-6165', 'Alex Rodriguez: (810) 219-7286', 'John Miller: (673) 805-8394', 'Laura Brown: (843) 583-8123', 'Sarah Wilson: (904) 357-9522', 'David Garcia: (639) 976-2570', 'Jane Wilson: (465) 394-2734', 'Michael Johnson: (385) 411-4435', 'Alex Miller: (444) 409-7942', 'Michael Williams: (224) 409-4153', 'Emily Johnson: (935) 435-8698', 'Chris Garcia: (274) 991-5495', 'Katie Miller: (460) 132-8816', 
```

```python

```