# Module 2 Assignment:

In this assignment, you will be implementing a CRF Model in python.


```python
import sklearn_crfsuite
from sklearn_crfsuite import metrics
import json
from sklearn.model_selection import train_test_split
import numpy as np

with open("assets/list_of_real_usa_addresses.json", "r") as file:
    addresses_json = json.load(file)
```

# Task 1

First, we must prepare the data. 

We have read in address_json, so in address_json, we have a list of dictionaries that looks like this:

`[{'address': '777 Brockton Avenue', 'city': 'Abington', 'state': 'MA', 'zip': '2351'}, 
{'address': '30 Memorial Drive', 'city': 'Avon', 'state': 'MA', 'zip': '2322'}, .... etc.]`

For your task_1 function, your job is to implement a function that creates X and y. X will be a list of lists of the tokens for addresses, and y will be the corresponding fields for X. 

To be more clear, taking the above example entries, we would want:

`X = [['777', 'Brockton', 'Avenue', 'Abington', 'MA', '2351'], ['30', 'Memorial', 'Drive', 'Avon', 'MA', '2322']]`

and

`y = [['address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state', 'zip']]`

You may find the `.split()` method for python strs to be helpful.

<!-- grade_id: cell-7e16173e433a44bf, points: 0, grade: true, solution: false -->
```python
task_id = '1'

## Utility Code will go here

```

<!-- grade_id: cell-4372d119ebe834d5, points: 0, grade: false, solution: true -->
```python
def task_1_solution():
    X = []
    y = []


    one_x = []
    one_y = []
    for i in addresses_json:

        for k in i.keys():
            
            one_x = one_x + i[k].split(' ')
            one_y = one_y + [k] * len(i[k].split(' '))

        X.append(one_x)
        y.append(one_y)

        one_x = []
        one_y = []
        

    return X, y
```

```python
# Use this cell to explore your solution.

task_1_solution()[0][0]
```

Output:
```
['777', 'Brockton', 'Avenue', 'Abington', 'MA', '2351']
```

```python
task_1_solution()[1][0]
```

Output:
```
['address', 'address', 'address', 'city', 'state', 'zip']
```

<!-- grade_id: cell-54ae0d067c9ce8b2, points: 20, grade: true, solution: false -->
```python
print(f"Task {task_id} - AG tests")

stu_X, stu_y = task_1_solution()

print(f"Task {task_id} - your X:\n{stu_X[:2]}")  # Display a subset of the data for clarity
print(f"Task {task_id} - your y:\n{stu_y[:2]}")  # Display a subset of the data for clarity

assert isinstance(stu_X, list), f"Task {task_id}: The returned `X` should be a list."
assert isinstance(stu_y, list), f"Task {task_id}: The returned `y` should be a list."

assert len(stu_X) == len(
    stu_y
), f"Task {task_id}: The lengths of `X` and `y` should match. Found len(X)={len(stu_X)} and len(y)={len(stu_y)}."

assert all(
    isinstance(seq, list) and all(isinstance(word, str) for word in seq)
    for seq in stu_X
), f"Task {task_id}: Each element of `X` should be a list of strings."

assert all(
    isinstance(seq, list) and all(isinstance(label, str) for label in seq)
    for seq in stu_y
), f"Task {task_id}: Each element of `y` should be a list of strings."


```

Output:
```
Task 1 - AG tests
Task 1 - your X:
[['777', 'Brockton', 'Avenue', 'Abington', 'MA', '2351'], ['30', 'Memorial', 'Drive', 'Avon', 'MA', '2322']]
Task 1 - your y:
[['address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state', 'zip']]
```

# Task 2

Next in CRF, we need to extract features for each token. A convenient way to do this is to write a function that creates a dictionary for a token containing information about the token. The function is started for you below. Add the keys 'is_digit', 'length', and 'is_upper' to the dictionary. 

<!-- grade_id: cell-b9ab500e5e6341bc, points: 0, grade: true, solution: false -->
```python
task_id = '2'

## Utility Code will go here

```

<!-- grade_id: cell-a21658df18118364, points: 0, grade: false, solution: true -->
```python
def task_2_solution(token):
    
    return {
        "word": token,   
        'is_digit': token.isdigit(), 
        'length': len(token),
        'is_upper':token.isupper()
    }
```

```python
# Use this cell to explore your solution.

task_2_solution("apple")
```

Output:
```
{'word': 'apple', 'is_digit': False, 'length': 5, 'is_upper': False}
```

<!-- grade_id: cell-1d570f3025d4abc3, points: 20, grade: true, solution: false -->
```python
print(f"Task {task_id} - AG tests")


teststr1 = "apple"

teststr2 = "77"

teststr3 = "MI"


stu_ans1 = task_2_solution(teststr1)

stu_ans2 = task_2_solution(teststr2)

stu_ans3 = task_2_solution(teststr3)

assert all(
    isinstance(entry, dict) for entry in [stu_ans1, stu_ans2, stu_ans3]
), (
    f"Task {task_id}: task_2_solution should return a dict "
)



```

Output:
```
Task 2 - AG tests
```

# Task 3

Now that you've crated a function to create a dictionary for a new token, write a function that takes the X list you created in task 1 and replaces each token with the dictionary returned by your task 2 function. 

<!-- grade_id: cell-1a0c38c3906669ed, points: 0, grade: true, solution: false -->
```python
task_id = '3'

## Utility Code will go here

```

<!-- grade_id: cell-08793c70677f13c6, points: 0, grade: false, solution: true -->
```python
def task_3_solution():
    old_X , _ = task_1_solution()

    X = []

    one_x = []

    for i in old_X:

        for k in i:

            one_x.append(task_2_solution(k))

        X.append(one_x)
        one_x = []
            
    
    return X
```

```python
# Use this cell to explore your solution.

task_3_solution()[0]
```

Output:
```
[{'word': '777', 'is_digit': True, 'length': 3, 'is_upper': False},
 {'word': 'Brockton', 'is_digit': False, 'length': 8, 'is_upper': False},
 {'word': 'Avenue', 'is_digit': False, 'length': 6, 'is_upper': False},
 {'word': 'Abington', 'is_digit': False, 'length': 8, 'is_upper': False},
 {'word': 'MA', 'is_digit': False, 'length': 2, 'is_upper': True},
 {'word': '2351', 'is_digit': True, 'length': 4, 'is_upper': False}]
```

<!-- grade_id: cell-51d761b69937207a, points: 20, grade: true, solution: false -->
```python
print(f"Task {task_id} - AG tests")
stu_ans = task_3_solution()

print(f"Task {task_id} - your answer:\n{stu_ans}")

assert isinstance(
    stu_ans, list
), f"Task {task_id}: Your function should return a list"


assert all(
    isinstance(entry, list) for entry in stu_ans
), (
    f"Task {task_id}: Your returned list should be a list of lists "
)


```

Output:
```
Task 3 - AG tests
Task 3 - your answer:
[[{'word': '777', 'is_digit': True, 'length': 3, 'is_upper': False}, {'word': 'Brockton', 'is_digit': False, 'length': 8, 'is_upper': False}, {'word': 'Avenue', 'is_digit': False, 'length': 6, 'is_upper': False}, {'word': 'Abington', 'is_digit': False, 'length': 8, 'is_upper': False}, {'word': 'MA', 'is_digit': False, 'length': 2, 'is_upper': True}, {'word': '2351', 'is_digit': True, 'length': 4, 'is_upper': False}], [{'word': '30', 'is_digit': True, 'length': 2, 'is_upper': False}, {'word': 'Memorial', 'is_digit': False, 'length': 8, 'is_upper': False}, {'word': 'Drive', 'is_digit': False, 'length': 5, 'is_upper': False}, {'word': 'Avon', 'is_digit': False, 'length': 4, 'is_upper': False}, {'word': 'MA', 'is_digit': False, 'length': 2, 'is_upper': True}, {'word': '2322', 'is_digit': True, 'length': 4, 'is_upper': False}], [{'word': '250', 'is_digit': True, 'length': 3, 'is_upper': False}, {'word': 'Hartford', 'is_digit': False, 'length': 8, 'is_
```

# Task 4

Now you have your y array and new X array, you are ready to train a CRF model. 

In the following function, create a crf model with algorithm='lbfgs', c1=0.1, c2=0.1, max_iterations=100, and all_possible_transitions=True. Make sure you use CRF model from `sklearn_crfsuite`. 

We have split the X and y arrays into train and test sets. Fit the crf model on the 'train' set with `crf.fit()`, and make a prediction on X_test with `crf.predict()`. 

Your function should return the predicted y array and the `y_test` variable we've defined for you.

Refer [here](https://sklearn-crfsuite.readthedocs.io/en/latest/tutorial.html) for a tutorial on how to accomplish this. You only need to pull 3 lines of code from the tutorial.

<!-- grade_id: cell-a7e17b71832bde2c, points: 0, grade: true, solution: false -->
```python
task_id = '4'

## Utility Code will go here


```

<!-- grade_id: cell-e00e72676c294cc5, points: 0, grade: false, solution: true -->
```python
def task_4_solution():

    _, y = task_1_solution()
    X = task_3_solution()

    # Split into train and test sets
    X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.2, random_state=42)


    crf = sklearn_crfsuite.CRF(algorithm='lbfgs', c1=0.1, c2=0.1, max_iterations=100,all_possible_transitions=True)
    crf.fit(X_train, y_train)

    y_pred = crf.predict(X_test)

    return y_pred, y_test



```

```python
# Use this cell to explore your solution.

task_4_solution()[0][0]
```

Output:
```
['address', 'address', 'address', 'address', 'city', 'state', 'zip']
```

<!-- grade_id: cell-c419a7ceb0561a74, points: 20, grade: true, solution: false -->
```python

print(f"Task {task_id} - AG tests")
stu_ans = task_4_solution()

print(f"Task {task_id} - your answer:\n{stu_ans}")


assert isinstance(
    stu_ans[0], list
), f"Task {task_id}: Your function should return a list. "

```

Output:
```
Task 4 - AG tests
Task 4 - your answer:
([['address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state', 'zip'], ['address', 'address', 'address', 'city', 'state
```

You have just utilized the CRF model to predict the field of every token in the X_test dataset!

# Task 5: Evaluate Results:

Compare the y_pred you got with y_test to see how accurately the crf predicted each label.

`metrics.flat_accuracy_score()` will be useful here.

<!-- grade_id: cell-56ca93fc5d66494b, points: 0, grade: true, solution: false -->
```python
task_id = '5'

## Utility Code will go here

```

<!-- grade_id: cell-585422c2320c19dd, points: 0, grade: false, solution: true -->
```python
def task_5_solution():

    y_pred, y_test = task_4_solution()

    accuracy = metrics.flat_accuracy_score(y_test, y_pred)

    

    return accuracy


```

```python
# Use this cell to explore your solution.

task_5_solution()
```

Output:
```
0.9684542586750788
```

<!-- grade_id: cell-212b339ca289fe33, points: 20, grade: true, solution: false -->
```python
print(f"Task {task_id} - AG tests")
stu_ans = task_5_solution()

print(f"Task {task_id} - your answer:\n{stu_ans}")

assert isinstance(
    stu_ans, float
), f"Task {task_id}: Your function should return a float. "



```

Output:
```
Task 5 - AG tests
Task 5 - your answer:
0.9684542586750788
```

```python

```